BEHAV3D Tumor Profiler to map heterogeneous cancer cell behavior in the tumor microenvironment

Author response:

The following is the authors’ response to the current reviews

Reviewer #1 (Public review):

In this work, Rios-Jimenez and Zomer et al have developed a 'zero-code' accessible computational framework (BEHAV3D-Tumour Profiler) designed to facilitate unbiased analysis of Intravital imaging (IVM) data to investigate tumour cell dynamics (via the tool's central 'heterogeneity module' ) and their interactions with the tumour microenvironment (via the 'large-scale phenotyping' and 'small-scale phenotyping' modules). A key strength is that it is designed as an open-source modular Jupyter Notebook with a user-friendly graphical user interface and can be implemented with Google Colab, facilitating efficient, cloud-based computational analysis at no cost. In addition, demo datasets are available on the authors GitHub repository to aid user training and enhance the usability of the developed pipeline.

To demonstrate the utility of BEHAV3D-TP, they apply the pipeline to timelapse IVM imaging datasets to investigate the in vivo migratory behaviour of fluorescently labelled DMG cells in tumour bearing mice. Using the tool's 'heterogeneity module' they were able to identify distinct single-cell behavioural patterns (based on multiple parameters such as directionality, speed, displacement, distance from tumour edge) which was used to group cells into distinct categories (e.g. retreating, invasive, static, erratic). They next applied the framework's 'large-scale phenotyping' and 'small-scale phenotyping' modules to investigate whether the tumour microenvironment (TME) may influence the distinct migratory behaviours identified. To achieve this, they combine TME visualisation in vivo during IVM (using fluorescent probes to label distinct TME components) or ex vivo after IVM (by large-scale imaging of harvested, immunostained tumours) to correlate different tumour behavioural patterns with the composition of the TME. They conclude that this tool has helped reveal links between TME composition (e.g. degree of vascularisation, presence of tumour-associated macrophages) and the invasiveness and directionality of tumour cells, which would have been challenging to identify when analysing single kinetic parameters in isolation.

While the analysis provides only preliminary evidence in support of the authors conclusions on DMG cell migratory behaviours and their relationship with components of the tumour microenvironment, conclusions are appropriately tempered in the absence of additional experiments and controls.

The authors also evaluated the BEHAV3D TP heterogeneity module using available IVM datasets of distinct breast cancer cell lines transplanted in vivo, as well as healthy mammary epithelial cells to test its usability in non-tumour contexts where the migratory phenotypes of cells may be more subtle. This generated data is consistent with that produced during the original studies, as well as providing some additional (albeit preliminary) insights above that previously reported. Collectively, this provides some confidence in BEHAV3D TP's ability to uncover complex, multi-parametric cellular behaviours that may be missed using traditional approaches.

While the tool does not facilitate the extraction of quantitative kinetic cellular parameters (e.g. speed, directionality, persistence and displacement) from intravital images, the authors have developed their tool to facilitate the integration of other data formats generated by open-source Fiji plugins (e.g. TrackMate, MTrackJ, ManualTracking) which will help ensure its accessibility to a broader range of researchers. Overall, this computational framework appears to represent a useful and comparatively user-friendly tool to analyse dynamic multi-parametric data to help identify patterns in cell migratory behaviours, and to assess whether these behaviours might be influenced by neighbouring cells and structures in their microenvironment.

When combined with other methods, it therefore has the potential to be a valuable addition to a researcher's IVM analysis 'tool-box'.

We thank the reviewer for carefully considering our manuscript and providing constructive comments. We appreciate the recognition of BEHAV3D-TP’s user-friendliness, modular design, and ability to link cell behavior with the tumor microenvironment. In the future, we plan to extend the tool to incorporate segmentation and tracking modules, once we have approaches that are broadly applicable or allow for personalized model training, further enhancing its utility for the community.

Reviewer #2 (Public review):

Summary:

The authors produce a new tool, BEHAV3D to analyse tracking data and to integrate these analyses with large and small scale architectural features of the tissue. This is similar to several other published methods to analyse spatio-temporal data, however, the connection to tissue features is a nice addition, as is the lack of requirement for coding. The tool is then used to analyse tracking data of tumour cells in diffuse midline glioma. They suggest 7 clusters exist within these tracks and that they differ spatially. They ultimately suggest that these behaviours occur in distinct spatial areas as determined by CytoMAP.

Strengths:

- The tool appears relatively user-friendly and is open source. The combination with CytoMAP represents a nice option for researchers.

- The identification of associations between cell track phenotype and spatial features is exciting and the diffuse midline glioma data nicely demonstrates how this could be used.

We thank the reviewer for their careful reading and thoughtful comments. Feedback from all revision rounds has helped us clarify key points and improve the manuscript, and we are grateful for the positive remarks regarding our application to diffuse midline glioma and the potential of the tool to enable new biological insights.

Reviewer #3 (Public review):

The manuscript by Rios-Jimenez developed a software tool, BEHAV3D Tumor Profiler, to analyze 3D intravital imaging data and identify distinctive tumor cell migratory phenotypes based on the quantified 3D image data. Moreover, the heterogeneity module in this software tool can correlate the different cell migration phenotypes with variable features of the tumor microenvironment. Overall, this is a useful tool for intravital imaging data analysis and its open-source nature makes it accessible to all interested users.

Strengths:

An open-source software tool that can quantify cell migratory dynamics from intravital imaging data and identify distinctive migratory phenotypes that correlate with variable features of the tumor microenvironment.

Weaknesses:

Motility is the main tumor cell feature analyzed in the study together with some other tumor-intrinsic features, such as morphology. However, these features are insufficient to characterize and identify the heterogeneity of the tumor cell population that impacts their behaviors in the complex tumor microenvironment (TME). For instance, there are important non-tumor cell types in the TME, and the interaction dynamics of tumor cells with other cell types, e.g., ﬁbroblasts and distinct immune cells, play a crucial role in regulating tumor behaviors. BEHAV3D-TP focuses on analysis of tumor-alone features, and cannot be applied to analyze important cell-cell interaction dynamics in 3D.

We thank the reviewer for their careful assessment and encouraging remarks regarding BEHAV3D-TP.

Regarding the concern about the tool’s current focus on motility features, we would like to clarify again that BEHAV3D-TP is designed to be highly flexible and extensible. Users can incorporate a wide range of features—including dynamic, morphological, and spatial parameters—into their analyses. In the latest revision, we have make this even more explicit by explaining that the feature selection interface allows users to either (i) directly select them for clustering or (ii) select features for correlation with clusters (See Small scale phenotyping module section in Methods).

Importantly, while our current analysis emphasizes clustering based on dynamic behaviors, Figure 4 demonstrates that these behavioral clusters are associated at the single-cell level with distinct proximities to key TME components, such as TAMMs and blood vessels. These spatial interaction features could also have been included in the clustering itself—creating dynamic-spatial clusters—but we deliberately chose not to do so. This decision was guided by established principles of feature selection: including features with unknown or potentially irrelevant variability can introduce noise and obscure biologically meaningful patterns, ultimately reducing the clarity and interpretability of the resulting clusters. Instead, we adopted a two-step approach—first identifying clusters based on core dynamic features, then examining their relationships with spatial and interaction metrics. This allowed us to reveal meaningful associations of particular cell behavior such as the invading cluster in proximity of TAMMs without overfitting or complicating the clustering model.

To address the reviewer’s point in the latest revision round, we have updated the Small-scale phenotyping module to highlight the possibility of including spatial interaction features with various TME cell types. We also revised the manuscript text and Figure 1 to clarify that these environmental features can be used both upstream as clustering input (Option 1) and for downstream analysis (Option 2), depending on the user’s experimental goals. Attached to this rebuttal letter, we also provide an additional figure illustrating these options in the feature selection panels of the Colab notebook.

In summary, while the clustering presented in this study is based on dynamic parameters, BEHAV3D-TP fully supports the integration of interaction features and other non-motility descriptors. This modularity enables users to customize their analysis pipelines according to specific biological questions, including those involving cell–cell interactions and spatial dynamics within the TME.

The following is the authors’ response to the original reviews.

Reviewer #1 (Public review):

Summary:

Intravital microscopy (IVM) is a powerful tool that facilitates live imaging of individual cells over time in vivo in their native 3D tissue environment. Extracting and analysing multi-parametric data from IVM images however is challenging, particularly for researchers with limited programming and image analysis skills. In this work, RiosJimenez and Zomer et al have developed a 'zero-code' accessible computational framework (BEHAV3D-Tumour Profiler) designed to facilitate unbiased analysis of IVM data to investigate tumour cell dynamics (via the tool's central 'heterogeneity module' ) and their interactions with the tumour microenvironment (via the 'large-scale phenotyping' and 'small-scale phenotyping' modules). It is designed as an open-source modular Jupyter Notebook with a user-friendly graphical user interface and can be implemented with Google Colab, facilitating efficient, cloud-based computational analysis at no cost. Demo datasets are also available on the authors GitHub repository to aid user training and enhance the usability of the developed pipeline.

To demonstrate the utility of BEHAV3D-TP, they apply the pipeline to timelapse IVM imaging datasets to investigate the in vivo migratory behaviour of fluorescently labelled DMG cells in tumour bearing mice. Using the tool's 'heterogeneity module' they were able to identify distinct single-cell behavioural patterns (based on multiple parameters such as directionality, speed, displacement, distance from tumour edge) which was used to group cells into distinct categories (e.g. retreating, invasive, static, erratic). They next applied the framework's 'large-scale phenotyping' and 'small-scale phenotyping' modules to investigate whether the tumour microenvironment (TME) may influence the distinct migratory behaviours identified. To achieve this, they combine TME visualisation in vivo during IVM (using fluorescent probes to label distinct TME components) or ex vivo after IVM (by large-scale imaging of harvested, immunostained tumours) to correlate different tumour behavioural patterns with the composition of the TME. They conclude that this tool has helped reveal links between TME composition (e.g. degree of vascularisation, presence of tumour-associated macrophages) and the invasiveness and directionality of tumour cells, which would have been challenging to identify when analysing single kinetic parameters in isolation.

The authors also evaluated the BEHAV3D TP heterogeneity module using available IVM datasets of distinct breast cancer cell lines transplanted in vivo, as well as healthy mammary epithelial cells to test its usability in non-tumour contexts where the migratory phenotypes of cells may be more subtle. This generated data is consistent with that produced during the original studies, as well as providing some additional (albeit preliminary) insights above that previously reported. Collectively, this provides some confidence in BEHAV3D TP's ability to uncover complex, multi-parametric cellular behaviours that may be missed using traditional approaches.

Overall, this computational framework appears to represent a useful and comparatively user-friendly tool to analyse dynamic multi-parametric data to help identify patterns in cell migratory behaviours, and to assess whether these behaviours might be influenced by neighbouring cells and structures in their microenvironment. When combined with other methods, it therefore has the potential to be a valuable addition to a researcher's IVM analysis 'tool-box'.

Strengths:

• Figures are clearly presented, and the manuscript is easy to follow.

• The pipeline appears to be intuitive and user-friendly for researchers with limited computational expertise. A detailed step-by-step video and demo datasets are also included to support its uptake.

• The different computational modules have been tested using relevant datasets, including imaging data of normal and tumour cells in vivo.

• All code is open source, and the pipeline can be implemented with Google Colab.

• The tool combines multiple dynamic parameters extracted from timelapse IVM images to identify single-cell behavioural patterns and to cluster cells into distinct groups sharing similar behaviours, and provides avenues to map these onto in vivo or ex vivo imaging data of the tumour microenvironment

Weaknesses:

• The tool does not facilitate the extraction of quantitative kinetic cellular parameters (e.g. speed, directionality, persistence and displacement) from intravital images. To use the tool researchers must first extract dynamic cellular parameters from their IVM datasets using other software including Imaris, which is expensive and therefore not available to all. Nonetheless, the authors have developed their tool to facilitate the integration of other data formats generated by open-source Fiji plugins (e.g. TrackMate, MTrackJ, ManualTracking) which will help ensure its accessibility to a broader range of researchers.

• The analysis provides only preliminary evidence in support of the authors conclusions on DMG cell migratory behaviours and their relationship with components of the tumour microenvironment. The authors acknowledge this however, and conclusions are appropriately tempered in the absence of additional experiments and controls.

We thank the reviewer for their thorough and constructive assessment of our work and are pleased that the accessibility, functionality, and potential impact of BEHAV3DTumour Profiler were well received. We particularly appreciate the acknowledgment of the tool’s ease of use for researchers with limited computational expertise, the clarity of the manuscript, and the relevance of our approach for identifying multi-parametric migratory behaviours and their correlation with the tumour microenvironment.

Regarding the weaknesses raised:

(1) Lack of built-in tracking and kinetic parameter extraction – As noted in our initial revision, while we agree that integrating open-source tracking and segmentation functionality could be valuable, it is beyond the scope of the current work. Our tool is designed to focus specifically on downstream analysis of already extracted kinetic data, addressing a gap in post-processing tools for exploring complex migratory behaviour and spatial correlations. Since different experimental systems often require tailored imaging and segmentation pipelines, we believe that decoupling tracking from the downstream analysis can actually be a strength, offering greater versatility. Researchers can use their preferred or most appropriate tracking software—whether proprietary or opensource—and then analyze the resulting data with BEHAV3D-TP. To support this, we ensured compatibility with widely used tools including open-source Fiji plugins (e.g., TrackMate, MTrackJ, ManualTracking), and we also cited several relevant studies and that address the upstream processing steps. Importantly, the main aim of our tool is to fill the gap in post-tracking analysis, enabling quantitative interpretation and pattern recognition that has until now required substantial coding effort or custom solutions.

(2) Preliminary nature of the biological conclusions – We fully agree with this assessment and have explicitly acknowledged this limitation in the manuscript. Our aim was to demonstrate the utility of BEHAV3D-TP in uncovering heterogeneity and spatial associations in vivo, while encouraging further hypothesis-driven studies using complementary biological approaches. We are grateful that the reviewer recognizes the cautious interpretation of our results and their added value beyond single-parameter analysis.

Reviewer #2 (Public review):

Summary:

The authors produce a new tool, BEHAV3D to analyse tracking data and to integrate these analyses with large and small scale architectural features of the tissue. This is similar to several other published methods to analyse spatio-temporal data, however, the connection to tissue features is a nice addition, as is the lack of requirement for coding. The tool is then used to analyse tracking data of tumour cells in diffuse midline glioma. They suggest 7 clusters exist within these tracks and that they differ spatially. They ultimately suggest that there these behaviours occur in distinct spatial areas as determined by CytoMAP.

Strengths:

- The tool appears relatively user-friendly and is open source. The combination with CytoMAP represents a nice option for researchers.

- The identification of associations between cell track phenotype and spatial features is exciting and the diffuse midline glioma data nicely demonstrates how this could be used.

Weaknesses:

The revision has dealt with many concerns, however, the statistics generated by the process are still flawed. While the statistics have been clarified within the legends and this is a great improvement in terms of clarity the underlying assumptions of the tests used are violated. The problem is that individual imaging positions or tracks are treated as independent and then analysed by ANOVA. As separate imaging positions within the same mouse are not independent, nor are individual cells within a single mouse, this makes the statistical analyses inappropriate. For a deeper analysis of this that is feasible within a review please see Lord, Samuel J., et al. "SuperPlots: Communicating reproducibility and variability in cell biology." The Journal of cell biology 219.6 (2020): e202001064. Ultimately, while this is a neat piece of software facilitating the analysis of complex data, the fact that it will produce flawed statistical analysis is a major problem. This problem is compounded by the fact that much imaging analysis has been analysed in this inappropriate manner in the past, leading to issues of interpretation and ultimately reproducibility.

We thank the reviewer for their careful reading and thoughtful feedback. We are encouraged by the recognition of BEHAV3D-TP’s ease of use, open-source accessibility, and the value of integrating cell behaviour with spatial features of the tissue. We appreciate the positive remarks regarding our application to diffuse midline glioma (DMG) and the potential for the tool to enable new biological insights.

We also appreciate the reviewer’s continued concern regarding the statistical treatment of the data. While we agree with the broader principle that care must be taken to avoid violating assumptions of independence, we respectfully disagree that all instances where individual tracks or imaging positions are used constitute flawed analysis. Importantly, our work is centered on characterizing heterogeneity at the single-cell level in distinct TME regions. Therefore, in certain cases—especially when comparing distinct behavioral subtypes across varying TME environments and multiple mice—it is appropriate to treat individual imaging positions as independent units. This approach is particularly relevant given our findings that large-scale TME regions differ across positions. When analyzing features such as the percentage of DMG cells in proximity to TAMMs, averaging per mouse would obscure these regional differences and reduce the resolution of biologically meaningful variation.

To address this concern further, we have revised the figure legends, main text, and documentation, carefully considering the appropriate statistical unit for each analysis. As detailed below, we used mouse-level aggregation where the experimental question required inter-mouse reproducibility, and a position-based approach where the aim was to explore intra-tumoral heterogeneity.

Figure 3d and Supplementary Figure 5d: In this analysis, we treated imaging positions as independent units because our data specifically demonstrate that, within individual mice, different positions correspond to distinct large-scale tumor microenvironment phenotypes. Therefore, averaging across the whole mouse would obscure these important spatial differences and not accurately reflect the heterogeneity we aim to characterize.

Figure 4c-e; Supplementary Figure 6d: While our initial aim was to highlight single-cell variability, we acknowledge that the original presentation may have been misleading. In the revised manuscript, we have updated the graphs for greater clarity. To quantify how often tumor cells of each behavioral type are located near TAMMs (Fig. 4c) or blood vessels (Fig. 4e), we now calculate the percentage of tumor cells "close" to environmental feature per behavioral cluster within each imaging position. This classification is based on the distance to the TME feature of interest and is detailed in the “Large-scale phenotyping” section of the Methods. For the number of SR101 objects in a 30um radius we averaged per position.

We treated individual imaging positions as the units of analysis rather than averaging per mouse, as our data (see Figure 2) show that positions vary in their TME phenotypes—such as Void, TAMM/Oligo, and TAMM/Vascularized—as well as in the number of TAMMs, SR101 cells or blood vessels per position. These differences are biologically meaningful and relevant to the quantification that we performed – percentage of tumor cell in close proximity to distinct TME features.

To account for inter-mouse and TME region variability, we applied a linear mixedeffects model with both mouse and TME class included as random effects.

Supplementary Figure 3d: Following the reviewer’s suggestion, we have averaged the distance to the 3 closest GBM neighbours per mouse, treating each mouse as an independent unit for comparison across distinct GBM morphodynamic clusters. To account for inter-mouse variability when assessing statistical significance, we employed a linear mixed model with mouse included as a random effect.

Distance to 3 neighbours is a feature not used in the clustering, thus variability between mice can be more pronounced—for example, due to differences in tumor compactness or microenvironment structure across individual mice. To appropriately account for this, mouse was included as a random effect in the model.

Supplementary Figure 4c: Following the reviewer’s suggestion, we averaged cell speed per mouse, treating each mouse as an independent unit for comparison across distinct DMG behavioral clusters. Statistical significance was assessed using ANOVA followed by Tukey’s post hoc test. When comparing cell speed, which is a feature used in the clustering process, inter-mouse variability was already addressed during clustering itself. Therefore, in the downstream analysis of this cluster-derived feature, it is appropriate to treat each mouse as an independent unit without including mouse as a random effect.

Supplementary Figure 5e-g: Following the reviewer’s suggestion, we averaged cell speed per mouse, treating each mouse as an independent unit for comparison across distinct DMG behavioral clusters. Statistical significance was assessed using ANOVA followed by Tukey’s post hoc test.

Supplementary Figure 6c: Following the reviewer’s suggestion, we averaged cell distance to the 10 closest DMG neighbours per mouse, treating each mouse as an independent unit for comparison across distinct DMG behavioral clusters. To account for inter-mouse variability, we used a linear mixed model with mouse included as a random effect.

Reviewer #3 (Public review):

The manuscript by Rios-Jimenez developed a software tool, BEHAV3D Tumor Profiler, to analyze 3D intravital imaging data and identify distinctive tumor cell migratory phenotypes based on the quantified 3D image data. Moreover, the heterogeneity module in this software tool can correlate the different cell migration phenotypes with variable features of the tumor microenvironment. Overall, this is a useful tool for intravital imaging data analysis and its open-source nature makes it accessible to all interested users.

Strengths:

An open-source software tool that can quantify cell migratory dynamics from intravital imaging data and identify distinctive migratory phenotypes that correlate with variable features of the tumor microenvironment.

Weaknesses:

Motility is only one tumor cell feature and is probably not sufficient to characterize and identify the heterogeneity of the tumor cell population that impacts their behaviors in the complex tumor microenvironment (TME). For instance, there are important nontumor cell types in the TME, and the interaction dynamics of tumor cells with other cell types, e.g., fibroblasts and distinct immune cells, play a crucial role in regulating tumor behaviors. BEHAV3D-TP focuses on only motility feature analysis, and cannot be applied to analyze other tumor cell dynamic features or cell-cell interaction dynamics.

Regarding the concern about the tool’s current focus on motility features, we would like to clarify that BEHAV3D-TP is designed to be highly flexible and extensible. As described in our first revision, users can incorporate a wide range of features—including dynamic, morphological, and spatial parameters—into their analyses. In the current revision, we have make this even more explicit by explaining that the feature selection interface allows users to either (i) directly select them for clustering or (ii) select features for correlation with clusters (See Small scale phenotyping module section in Methods and Rebuttal Figure).

Importantly, while our current analysis emphasizes clustering based on dynamic behaviors, Figure 4 demonstrates that these behavioral clusters are associated at the single-cell level with distinct proximities to key TME components, such as TAMMs and blood vessels. These spatial interaction features could also have been included in the clustering itself—creating dynamic-spatial clusters—but we deliberately chose not to do so. This decision was guided by established principles of feature selection: including features with unknown or potentially irrelevant variability can introduce noise and obscure biologically meaningful patterns, ultimately reducing the clarity and interpretability of the resulting clusters. Instead, we adopted a two-step approach—first identifying clusters based on core dynamic features, then examining their relationships with spatial and interaction metrics. This allowed us to reveal meaningful associations of particular cell behavior such as the invading cluster in proximity of TAMMs without overfitting or complicating the clustering model.

To further address the reviewer’s point, we have updated the Small-scale phenotyping module to highlight the possibility of including spatial interaction features with various TME cell types. We also revised the manuscript text and Figure 1 to clarify that these environmental features can be used both upstream as clustering input (Option 1) and for downstream analysis (Option 2), depending on the user’s experimental goals. Author response image 1 illustrates these options in the feature selection panels of the Colab notebook.

(a) In the small-scale phenotyping module, microenvironmental factors (MEFs) detected in the segmented IVM movies are identified and their coordinates imported.

From here, there are two options: (b) include the relationship to these MEFs as a feature for clustering, or (c) exclude this relationship and instead correlate MEFs with cell behavior to assess potential spatial associations.

In summary, while the clustering presented in this study is based on dynamic parameters, BEHAV3D-TP fully supports the integration of interaction features and other non-motility descriptors. This modularity enables users to customize their analysis pipelines according to specific biological questions, including those involving cell–cell interactions and spatial dynamics within the TME.

Reviewer #2 (Recommendations for the authors):

If the software were adjusted to produce analyses following best practices in the field as outlined in Lord, Samuel J., et al. "SuperPlots: Communicating reproducibility and variability in cell biology." The Journal of cell biology 219.6 (2020): e202001064. this could be a helpful piece of software. The major current issue would be that it democratises the ability to analyse complex imaging data, allowing non-experts to carry out these analyses but misleads them and encourages poor statistical practice.

We appreciate the reviewer’s suggestion and the reference to best practices outlined in Lord et al., 2020. As discussed in detail in our point-by-point response to Reviewer #2, we have revised several figures to enhance clarity and statistical rigor, including Figure 4c,e; Supplementary Figures 3d, 4c, 5e–g, and 6c–d. Specifically, we adjusted how data are summarized and displayed—averaging per mouse where appropriate and clarifying the statistical methods used. Where imaging positions were retained as the unit of analysis, this decision was grounded in the biological relevance of intra-mouse spatial heterogeneity (as demonstrated in Figure 2). Additionally, we applied linear mixed-effects models in cases where inter-mouse or inter-Large scale TME regions variability needed to be accounted for. We believe these changes address the core concern about reproducibility and statistical interpretation while preserving the biological insights captured by our approach.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Intravital microscopy (IVM) enables live imaging of animals at single-cell level, offering essential insights into cancer progression. This technique allows for the observation of single-cell behaviors within their natural 3D tissue environments, shedding light on how genetic and microenvironmental changes influence the complex dynamics of tumors. IVM generates highly complex datasets that often exceed the analytical capacity of traditional uni-parametric approaches, which can neglect single-cell heterogeneous in vivo behavior and limit insights into microenvironmental influences on cellular behavior. To overcome these limitations, we present BEHAV3D Tumor Profiler (BEHAV3D-TP), a computational framework that enables unbiased single-cell classification based on a range of morphological and dynamic single cell features. BEHAV3D-TP integrates with widely used 2D and 3D image processing pipelines, enabling researchers without advanced computational expertise to profile cancer and healthy cell dynamics in IVM data. Here, we apply BEHAV3D-TP to study diffuse midline glioma (DMG), a highly aggressive pediatric brain tumor characterized by invasive progression. By extending BEHAV3D-TP to incorporate tumor microenvironment (TME) data from IVM or fixed correlative imaging, we demonstrate that distinct migratory behaviors of DMG cells are associated with specific TME components, including tumor-associated macrophages and vasculature. BEHAV3D-TP enhances the accessibility of computational tools for analyzing the complex behaviors of cancer cells and their interactions with the TME in IVM data.

Revised version: Intravital microscopy (IVM) enables live imaging of animals at single-cell level, offering essential insights into cancer progression. This technique allows for the observation of single-cell behaviors within their natural 3D tissue environments, shedding light on how genetic and microenvironmental changes influence the complex dynamics of tumors. IVM generates highly complex datasets that often exceed the analytical capacity of traditional uni-parametric approaches, which can neglect single-cell heterogeneous in vivo behavior and limit insights into microenvironmental influences on cellular behavior. To overcome these limitations, we present BEHAV3D Tumor Profiler (BEHAV3D-TP), a computational framework that enables unbiased single-cell classification based on a range of morphological, environmental and dynamic single cell features. BEHAV3D-TP integrates with widely used 2D and 3D image processing pipelines, enabling researchers without advanced computational expertise to profile cancer and healthy cell dynamics in IVM data. Here, we apply BEHAV3D-TP to study diffuse midline glioma (DMG), a highly aggressive pediatric brain tumor characterized by invasive progression. By extending BEHAV3D-TP to incorporate tumor microenvironment (TME) data from IVM or fixed correlative imaging, we demonstrate that distinct migratory behaviors of DMG cells are associated with specific TME components, including tumor-associated macrophages and vasculature. BEHAV3D-TP enhances the accessibility of computational tools for analyzing the complex behaviors of cancer cells and their interactions with the TME in IVM data.

Introduction

Cancer is a complex disease driven by a sequence of genetic and microenvironmental changes. As a result, tumors comprise many different cell populations that are in constant evolution1. This tumor heterogeneity has important biological and clinical implications as it differentially promotes key characteristics of individual cancer cells, including proliferation, invasion, and drug resistance 2. Yet, we are far from fully understanding all the dynamic and bi-directional interactions in heterogeneous tumors, and how these interactions influence cancer cell behavior at the single-cell level. Emerging single-cell technologies and analysis tools provide a new opportunity to profile individual cells within tumors. For example, recent spatial transcriptomic and proteomic profiling provides evidence for the existence of multiple distinct microenvironmental niches within one tumor 3. However, these techniques can only provide static information and make use of thin tissue sections that are only one cell layer thick. To investigate the real-time behavior of individual cells intravital microscopy (IVM) has been developed to study tumor cell behavior evolution over time but also in their three-dimensional native tissue environment.

Earlier version: IVM is a technique used to capture both spatial and temporal information at a single-cell level in living animals. Over the past years, it has provided many new insights into the dynamics of tumor heterogeneity and the development of therapy resistance 4–15. For example, intravital imaging of mouse breast cancer models revealed that only a small portion of tumor cells is motile 15, and that fast- and slow-migrating cells are present in distinct regions of the tumor 7,16. Moreover, additional heterogeneity exists among migrating tumor cells that can, for instance, move randomly or directionally relative to other tumor cells and/or surrounding healthy tissue 7. However, traditional IVM analyses have considered cancer cells or the tumor microenvironment (TME) as a homogeneous population, thereby ignoring heterogeneity at the single-cell level and the influence of the local microenvironment, where interactions between neighboring cells may affect cellular behavior. To address these complexities, recent studies have used unbiased approaches using morpho-kinetic parameters derived from live in vitro17 or in vivo imaging to classify cells and explore the single-cell behavioral landscape as an additional omic layer 18–20. In line with this concept of characterizing cellular dynamic properties for cell classification, we have previously developed an analytical platform termed BEHAV3D 19,21 allowing to perform behavioral phenotyping of engineered T cells targeting cancer. While BEHAV3D was initially developed to analyze T cell migratory behavior under controlled in vitro conditions, we sought to expand its application to investigate tumor cell behaviors in IVM data, where the complexity of the TME presents distinct analytical challenges. This manuscript builds on our foundational work but represents a significant advancement by adapting the pipeline specifically for IVM datasets.

Revised version: IVM is a technique used to capture both spatial and temporal information at a single-cell level in living animals. Over the past years, it has provided many new insights into the dynamics of tumor heterogeneity and the development of therapy resistance 4–15. For example, intravital imaging of mouse breast cancer models revealed that only a small portion of tumor cells is motile 15, and that fast- and slow-migrating cells are present in distinct regions of the tumor 7,16. Moreover, additional heterogeneity exists among migrating tumor cells that can, for instance, move randomly or directionally relative to other tumor cells and/or surrounding healthy tissue7. However, traditional IVM analyses have considered cancer cells or the tumor microenvironment (TME) as a homogeneous population, thereby ignoring heterogeneity at the single-cell level and the influence of the local microenvironment, where interactions between neighboring cells may affect cellular behavior. To address these complexities, recent studies have used unbiased approaches using morpho-kinetic parameters derived from live in vitro17 or in vivo imaging to classify cells and explore the single-cell behavioral landscape as an additional omic layer 18–20. In line with this concept of characterizing cellular dynamic properties for cell classification, we have previously developed an analytical platform termed BEHAV3D 19,21 allowing to perform behavioral phenotyping of engineered T cells targeting cancer. While BEHAV3D was initially developed to analyze T cell migratory behavior under controlled in vitro conditions, we sought to expand its application to investigate tumor cell behaviors in IVM data, where the complexity of the TME presents distinct analytical challenges. This manuscript builds on our foundational work but represents a significant advancement by adapting the pipeline specifically for IVM datasets.

While significant efforts have been made to develop open-source segmentation and tracking tools for live imaging data, including IVM22–27 fewer tools exist for the unbiased analysis of tumor dynamics. One major barrier is that implementing such analytical methods often requires substantial computational expertise, limiting accessibility for many biomedical researchers conducting IVM experiments. To bridge this gap, we present BEHAV3D Tumor Profiler (BEHAV3D-TP) by providing a robust, user-friendly tool that allows researchers to extract meaningful insights from dynamic cellular behaviors without requiring advanced programming skills. BEHAV3D-TP expands the capabilities of BEHAV3D by incorporating key features designed for in vivo tumor cell dynamics (Fig. 1). The pipeline includes extensions that integrate tumor cell behavior with spatial features within the native TME, using either in situ (IVM) or fixed correlative imaging to reveal how distinct spatial niches influence tumor dynamics.

BEHAV3D Tumor Profiler pipeline

Earlier version: Schematic representation of the workflow showing data preparation, input and outputs in each of the three modules. a, Time-lapse image processing and preparation, based on segmentation and tracking of tumor (and labelled TME cells). Mounting Google Drive is optional but recommended to work in the Google Colab framework. b, Heterogeneity Module provides distinct behavioral clusters for cells, and relates them to dynamic features and provides a back-projection of the tracked cells. c, d, Optional modules, additional to the Heterogeneity module. c, Large-scale phenotyping module combines TME large-scale information with cell morphodynamic profiling to further depict population frequencies in distinct TME regions. d, Small-scale phenotyping module combines TME small-scale information with morphodynamic profiling to depict how the cellular environment can influence tumor cell behavior.

Revised version: Schematic representation of the workflow showing data preparation, input and outputs in each of the three modules. a, Time-lapse image processing and preparation, based on segmentation and tracking of tumor (and labelled TME cells). Mounting Google Drive is optional but recommended to work in the Google Colab framework. b, Heterogeneity Module provides distinct behavioral clusters for cells, and relates them to dynamic features and provides a back-projection of the tracked cells. c, d, Optional modules, additional to the Heterogeneity module. c, Large-scale phenotyping module combines TME large-scale information with cell morphodynamic profiling to further depict population frequencies in distinct TME regions. d, Small-scale phenotyping module quantifies TME proximity and interaction features at the single-cell level. These features can be incorporated upstream as part of the clustering process to define behaviorally distinct subpopulations (option 1), or downstream to analyze how spatial relationships to TME components such as TAMMs or vasculature vary across behavioral clusters identified in the heterogeneity module (option 2).

To further facilitate unbiased analysis of tumor dynamics, we developed BEHAV3D-TP as a flexible tool that supports both commercial and open-source data formats, ensuring broad compatibility across different research workflows. The pipeline supports 2D and 3D data formats generated by Imaris and by Fiji plugins such as TrackMate, MTrackJ, and ManualTracking, enabling processing of datasets from the most widely used segmentation and tracking solutions. We demonstrate the versatility and broad applicability of BEHAV3D-TP by applying it to a range of published tumor and healthy tissue IVM datasets.

Next, we focused on Diffuse Midline Glioma (DMG), a highly aggressive pediatric brain tumor characterized by invasive growth and its capacity to spread to other regions of the brain and the spinal cord 28. Our approach allowed us to observe associations between DMG cell migration behavior and specific microenvironmental factors, including tumor-associated macrophages and vasculature, by visualizing TME components either during or after IVM. In our study, we specifically focused on cancer cell migration as a specific feature of DMG, but our methodology could easily be extended to other cancer cell behaviors such as intravasation or metastatic colonization. Overall, BEHAV3D-TP provides an accessible computational approach for a better understanding of tumor-TME crosstalk and the functional implications of this communication in heterogeneous tumors.

Results

BEHAV3D Tumor Profiler: an accessible tool for mapping heterogeneity in tumor dynamics

To comprehensively map and analyze the dynamic heterogeneity of tumor cells within their TME we developed BEHAV3D-TP [https://github.com/imAIgene-Dream3D/BEHAV3D_Tumor_Profiler]. To enhance accessibility and encourage widespread adoption among researchers studying single tumor cell dynamics in live imaging and inspired by previous work on democratizing advanced image analysis26,29,30, we designed BEHAV3D-TP as a modular Jupyter Notebook featuring an intuitive graphical user interface (GUI) (Figure 1 and Supplementary Video 1). The platform runs directly in a web browser and requires only a Google account for access. BEHAV3D-TP implementation within the Google Colaboratory (Colab) framework not only ensures efficient cloud-based computation with free and scalable resources but also enhances tool accessibility by eliminating the need for local software installation, a common barrier for inexperienced users. Regardless of programming proficiency, users can navigate the computational pipeline to extract comprehensive insights into tumor dynamics (Figure 1). After time-lapse image file processing and preparation using various frameworks (Figure 1a), the central component of BEHAV3D-TP — called the heterogeneity module (Figure 1b) — enables researchers to uncover diverse patterns of tumor morphodynamic behavior. Additionally, we provide two extension modules—the large-scale phenotyping module (Figure 1c) and the small-scale phenotyping module (Figure 1d)— which integrate tumor cell behavior identified with the heterogeneity module with TME features based on different data types. The large-scale phenotyping module uses post-IVM fixed imaging data for tissue phenotyping, identifying extensive regions with specific TME characteristics, and assessing the distribution of distinct tumor cell behavioral patterns within these regions. The small-scale phenotyping module uses in situ-labelled TME structures from IVM data, facilitating the correlation of these TME features within the microenvironment of individual tumor cell neighborhoods. Both modules, in conjunction with the heterogeneity module, help to identify how different aspects of the TME influence distinct behavioral patterns of tumor cells.

To evaluate the capability of the BEHAV3D-TP heterogeneity module in identifying diverse morphodynamic cell patterns, we applied it to IVM datasets from various tumor and healthy models. The datasets, processed with different segmentation and tracking tools, were analyzed for multivariate similarities in morphological and dynamic features (summarized in Table 1) using the dynamic time warping algorithm21.

As an initial benchmark, we analyzed a previously published IVM dataset of two migratory and metastatic breast cancer cell lines (4T1 and MDA-MB231), tracked in 2D using the MTrackJ Fiji plugin (Supplementary Figure 1a) 15. BEHAV3D-TP provided a more granular characterization of these behaviors, uncovering additional migratory patterns not reported in the original study (Supplementary Figure 1b-c). Specifically, we identified behaviors such as Fast, Intermediate, Very slow and Static that had a different distribution among both tumor cell lines (Supplementary Figure 1d-e).

To further assess the versatility of BEHAV3D-TP, we applied it to 2D IVM data of terminal end buds (TEB) from healthy mammary gland tissue31. This dataset features confetti-labeled epithelial cell types and was analyzed using maximum intensity projections of multiple z planes showing cell morphology changes. As described in the original publication31 cell segmentation was performed with Cellpose 2.032 within the MaSCOT-AI workflow, followed by tracking using the TrackMate plugin in Fiji. This workflow not only confirmed the compatibility of BEHAV3D-TP with advanced segmentation frameworks but also its usage to unravel distinct morphodynamic patterns in non-tumor epithelial cells (Supplementary Figure 2b,c). Consistent with previous findings that HR+ TEB (Pr-cre/Confetti) exhibit greater migratory capacity than HR-TEB (Elf5-rtTA/TetO-Cre/Confetti), our analysis suggests that this difference is primarily driven by a distinct morphodynamic behavior (Supplementary Figure 2 c,d). Specifically, HR+ TEB more frequently exhibited characteristics of Cluster 2, defined by a rounder shape and increased motility. Despite the small sample size requiring further validation, these results demonstrate BEHAV3D-TP’s ability to uncover complex, multiparametric cellular behaviors that are often missed in analyses focused on single-value comparisons across conditions.

We next tested BEHAV3D-TP with a 3D image processing pipeline applied to IVM data from adult glioma (GBM) 7. In this case, cells were segmented using μSAM33, a deep-learning foundation model adapted for microscopy data (Segment Anything for Microscopy), and subsequently tracked with TrackMate27 Fiji plugin (Supplementary Figure 3a). BEHAV3D-TP identified distinct GBM morphodynamic patterns, each characterized by unique combinations of morphological and dynamic features and associated with specific spatial distributions (Supplementary Figure 3b-d). Cluster 1, consisting of large and slow-moving cells, was more isolated and located farther from neighboring cells compared to Cluster 6 (small, fast, persistent cells) and Cluster 4 (very small, static cells). These findings suggest that the compact spatial arrangement of small cells within the tumor may contribute to their reduced size.

Collectively, these findings underscore the versatility and broad applicability of BEHAV3D-TP across diverse biological systems, dimensionalities (2D and 3D), imaging modalities, and data processing pipelines, supporting its utility as a robust tool for the in-depth analysis of cellular behavioral patterns in IVM data.

Identification of highly heterogeneous in vivo cell behavior displayed by DMGs

We then applied our platform to get insights into the behavioral heterogeneity among pediatric DMG tumors, characterized by a highly invasive nature. We implanted mice with patient derived DMG cells expressing H2BmNeonGreen marker and performed IVM for up to 5.4 hours upon tumor development (Figure supplement 4a). We tracked intravitally imaged DMG cells from different mice using Imaris, a commercial software a widely adapted by the IVM community34–40 for its intuitive 3D visualization and analysis capabilities. In each imaged position, we evaluated not only the dynamic characteristics of individual DMG cells but also their spatial relationship to the tumor edge, which was delineated using Imaris Surface module (Figure 2a, see Methods). Next, to untangle the complexity of tumor cell dynamics in an unbiased manner, we used the BEHAV3D-TP heterogeneity module, implementing multiparametric single-cell time-series classification to identify distinct single-cell behavioral patterns (Figure 2 a, b). For this, we used kinetic parameters, including displacement, speed, persistence, and movement direction relative to the tumor edge (Table 1), followed by dimensionality reduction to classify cells. The BEHAV3D-TP heterogeneity module identified 7 DMG behavioral clusters: retreating, slow retreating, erratic, static, slow, slow invading, and invading (Figure 2b, c, d, Supplementary video 2). To accurately describe each of the clusters, we analyzed their most predominant features (Figure 2d) and back-projected the cluster information into the original time-lapse to visually inspect cell behavior in relation to the tumor (Figure 2a, c, Figure supplement 4b, Supplementary video 2). Of particular interest were cells from the invading and retreating clusters, characterized by fast movement away from and towards the tumor edge, respectively (Figure 2c-f, and Figure supplement 4c). In fact, in this particular DMG model, approximately 10% of the cancer cells were displaying invasive behavior (Figure supplement 4d), in line with the intrinsic infiltrative nature of DMG41. Interestingly, about another 10% of DMG cells were moving back to the tumor edge (retreating, slow retreating clusters) (Figure supplement 4d), behavior that has been previously observed for glioblastoma and suggested to be regulated by differential chemotactic signaling by the tumor edge and the surrounding brain parenchyma7. We also observed erratic cells that display fast non-directed migration behavior, and static cells that are non-migratory (Figure 2c-f). Crucially, the distinct behavioral patterns were detected in all intravitally imaged mice (Figure supplement 4d), ruling out any sample bias that might lead to clusters consisting solely of cells from a single tumor or mouse. This underscores the effectiveness of the BEHAV3D-TP in revealing representative heterogeneous tumor behavioral patterns.

Behavioral profiling of DMG cells with BEHAV3D-TP heterogeneity module.

a, Representative intravitally imaged position showing DMG cells tracked relative to tumor Edge (left panel). Tracks were classified according to DMG behavior (right, color coded by cluster). Representative of n = 18 independent positions from n=6 mice. b, UMAP plot showing seven color-coded DMG cell behavioral clusters identified by BEHAV3D-TP. Each datapoint represents one DMG cell track. A total of 981 individual cells were tracked in n = 18 independent positions from n=6 mice. c, Representative intravitally imaged positions and enlarged sections for showcasing distinct DMG behavioral clusters Invading (yellow), Retreating (green), Erratic (dark purple) and Static (light purple) through the time series. n = 18 independent positions from n=6 mice. Scale bars 10µm. d, Heatmap depicting relative values of DMG cell features indicated for each cluster, named according to their most distinct characteristics. Arbitrary units in respect to maximal and minimal values for each feature. See Table 1 for features description. Bubble size according to significance levels, all features represented have a significance level of P < 0.001 (***): mean_displ2, mean squared displacement; speed; delta displacement; displacement length; persistence; invasion. P-values represent the significance of differences across clusters tested through ANOVA for each feature. e, f, UMAP representation of DMG cells’ velocity relative to the tumor edge(µm) (e) and speed (µm/h) (f). Each datapoint represents one cell track and is colored following a color gradient. For each feature colors represent the mean DMG track values. Representative of n = 18 independent positions from n=6 mice.

Mapping in vivo tumor cell migratory properties to distinct TME regions identified through ex vivo large-scale 3D imaging

To investigate whether the TME influences behavioral heterogeneity among cancer cells, we developed two complementary approaches to correlate the identified behavioral clusters (Figure 2) to the composition of the TME (Figures 1c, 3a and 4a).

Mapping distinct DMG behavioral patterns to distinct TME regions.

a, Overview of the large-scale TME region phenotyping workflow. After IVM imaging session, the tissue is fixed and stained and then analyzed using mLSR-3D. CytoMAP Spatial Analysis Toolbox is used to identify distinct large-scale regions, that are subsequently matched to IVM imaging dataset (c). b, Representative fixed multispectral image of a mLSR3D imaged DMG tumor and zoomed-in images of TME large-scale regions: Void (blue outline), TAMM/vascularized (red outline) and TAMM/Oligo (yellow outline). DMG nuclei (blue), CD31 (pink), Olig2 (purple), and Iba1 (yellow) are represented on the imaged. Scale bars, 100 µm (overview), 30 µm (zoomed-in regions). c, 3D multispectral image DMG cell rendering color-coded per large TME region: Void (blue), TAMM/Oligo (yellow) and TAMM/vascularized (red) (left). Based on the spatial coordinates of DMG cells classified into specific TME regions by CytoMAP, cells within the mLSR-3D images were subsequently assigned to their respective TME regions. Regions were then visually mapped on intravital imaging data (right). Representative of n=6 mice. White outline indicates one of the tumor regions analyzed with BEHAV3D-TP from the full tumor volume. d Boxplots showing the frequencies of different DMG behavioral clusters in Void, TAMM/Oligo and TAMM/vascularized TME large-scale regions. For each behavioral cluster, the y-axis displays the z-scored percentage of cells, normalized per mouse to account for inter-mouse variability. Each point represents an individual imaged position, with shape indicating the mouse of origin. P-values represent the significance of differences across TME regions tested through ANOVA with Tukey post hoc, for each behavioral cluster. n = 18 independent positions from n=6 mice.

The BEHAV3D-TP small-scale phenotyping correlates DMG behavioral profiles with microenvironmental composition at a single-cell resolution.

Earlier version: a, Overview of the small-scale TME phenotyping module. During IVM imaging, information about spatial distribution of SR101+, CD20r+, and CD31+ cells is collected and processed by the BEHAV3D-TP small-scale phenotyping module. Single DMG cell spatial TME features like distance to its neighbors, number of cells in a radius or minimal distance to a cell type were measured. b, Heatmap depicting DMG cell small-scale niches’ features across distinct DMG behavioral clusters. Arbitrary units in respect to maximal and minimal values for each feature. Bubble size according to significance levels, from smallest (P < 0.05 (**)) to medium (P < 0.01 (**)) to biggest (P < 0.001 (***)): dist_10_cell, Distance to nearest 10 cells, P < 0.01 (**); n_SR101, number of SR101 cells, P < 0.05 (*); min_SR101, minimal distance to a SR101 cell, P < 0.05 (*); n_CD20r, number of CD20r cells, P < 0.05 (*); min_Cd20r, minimal distance to a CD20r cell, P < 0.01 (**); min_BV_dist, minimal distance to a Blood vessel, P < 0.001 (***). P-values represent the significance of differences across clusters tested through ANOVA for each feature. c, e, Boxplot depicting the Minimal distance to a CD20r+ TAMM (µm) (c) and the Minimal distance to a blood vessel (µm) (e) across DMG behavioral clusters. Each point represents an individual tracked cell, color depicts the imaged position and shape denotes the mouse. P-values represent the significance of differences across behavioral clusters tested through ANOVA with Tukey’s post hoc. P < 0.01 (**);P < 0.05 (*).P < 0.001 (***); P < 0.1 (·). c, n = 10 independent positions from n=3 mice. e, n = 8 independent positions from n=3 mice. d, Multispectral image showing the movement of DMG cells (grey) relative to CD20r+ TAMM (blue) and the Tumor core. The trajectories show a path from the initial (blue) to the final position (red). Scale bar, 10 µm. f, Time lapse multispectral images (left to right) showing the DMG cells (blue) displacement along CD31 blood vessel (purple) away from the tumor edge. Scale bars, 30 µm.

Revised version: a, Overview of the small-scale TME phenotyping module. During IVM imaging, information about spatial distribution of SR101+, CD20r+, and CD31+ cells is collected and processed by the BEHAV3D-TP small-scale phenotyping module. Single DMG cell spatial TME features like distance to its neighbors, number of cells in a radius or minimal distance to a cell type were measured. b, Heatmap depicting DMG cell small-scale niches’ features across distinct DMG behavioral clusters. Arbitrary units in respect to maximal and minimal values for each feature. Bubble size according to significance levels, from smallest (P < 0.05 (**)) to medium (P < 0.01 (**)) to biggest (P < 0.001 (***)): dist_10_cell, Distance to nearest 10 cells, P < 0.01 (**); n_SR101, number of SR101 cells, P < 0.05 (*); min_SR101, minimal distance to a SR101 cell, P < 0.05 (*); n_CD20r, number of CD20r cells, P < 0.05 (*); min_Cd20r, minimal distance to a CD20r cell, P < 0.01 (**); min_BV_dist, minimal distance to a Blood vessel, P < 0.001 (***). P-values represent the significance of differences across clusters tested through ANOVA for each feature. c, e, Boxplots depicting the percentage of DMG cells across behavioral clusters that are close to CD20r+ TAMM (c) or to blood vessels (e). Each point represents an individual imaged position, color depicts the TME large-scale region type and shape denotes the mouse. Statistical differences across clusters were assessed using a linear mixed-effects model including mouse and TME class as random effects. P < 0.05 (*), P < 0.01 (**), P < 0.001 (***) indicate significant differences between clusters. c, n = 10 independent positions from n=3 mice. e, n = 8 independent positions from n=3 mice. d, Multispectral image showing the movement of DMG cells (grey) relative to CD20r+ TAMM (blue) and the Tumor core. The trajectories show a path from the initial (blue) to the final position (red). Scale bar, 10 µm. f, Time lapse multispectral images (left to right) showing the DMG cells (blue) displacement along CD31 blood vessel (purple) away from the tumor edge. Scale bars, 30 µm.

The first approach consists of performing large-scale TME phenotyping and identifying regions with a specific cellular composition and architecture within the TME of intravitally imaged tumors that were subsequently fixed. This is accomplished by using multispectral large-scale single-cell resolution 3D (mLSR-3D) imaging data42,43 and a spatial analysis framework called CytoMAP44 (Figure 3a). We selected microenvironmental components of interest, such as perivascular niches and glial cells, as they have been previously reported to play an important role in the invasive behavior of glioblastoma cells7,45–49. Specifically, intravitally imaged brains were collected immediately after the IVM imaging sessions and were fixed and stained for blood vessels (CD31), oligodendrocytes (Olig2), and tumor-associated microglia/macrophages (TAMM) (Iba1) (Figure 3b). Subsequent CytoMAP44 spatial phenotyping analysis identified three different environmental niches: Void regions (less vascularized and with low glial infiltration), TAMM/oligo regions (enriched in oligodendrocytes, TAMM, and tumor cells), and TAMM/vascularized regions (highly vascularized regions enriched in TAMM) (Figure 3b, and Figure supplement 5a-c). To investigate the heterogeneity of tumor cell behavior in relation to the identified microenvironmental niches, BEHAV3D-TP large-scale phenotyping module mapped the different niches (Void, TAMM/Oligo, and TAMM/vascularized) identified by CytoMAP onto 3D intravital imaging data (Figures 1c, 3c). In the TME-defined IVM-imaged positions we compared the frequencies of the different behavioral clusters that we identified with the BEHAV3D-TP heterogeneity module (Figure 3d). Interestingly, invading cells were more abundant in TAMM/vascularized regions compared to Void regions that contain a higher proportion of static cells, and TAMM/Oligo microenvironments contain more slow retreating cells compared to the other regions (Figure 3d, Figure supplement 5d). Finally, we compared the results obtained with BEHAV3D-TP to a more classical IVM analysis approach that relies on assessing single dynamic parameters. Interestingly, single parameters showed a restricted ability to identify differences in cellular behavior among various environments (Figure supplement 5e-g), underscoring the importance of multiparametric behavioral mapping in revealing more nuanced cellular behavior that cannot be captured solely by individual dynamic parameters.

Linkage of in vivo tumor cell behavior to the composition of the microenvironment at a single-cell level

Added in the revised version: With the goal of further refining TME phenotyping to better understand tumor cell behavior determinants, we implemented an alternative approach using the BEHAV3D-TP small-scale phenotyping module (Figure 1d, 4a). This module utilizes the detection of TME components during IVM to offer insights into their abundance and spatial relationship with DMG cells at a single-cell level. It complements the heterogeneity module, which captures tumor cell behavioral profiles, and offers two options of integration: (1) upstream integration, where TME-derived spatial features—such as proximity or interactions with specific components— are used as input variables for clustering tumor cells; and (2) downstream integration, where previously defined behavioral clusters are correlated with TME spatial features in a subsequent analytical step (Figures 1d, 4a). When using the upstream integration approach, it is essential to include only biologically meaningful features, as the presence of irrelevant or redundant variables can introduce noise and compromise the interpretability of the resulting clusters.

Earlier version: With the goal of further refining TME phenotyping to better understand tumor cell behavior determinants, we implemented an alternative approach using the BEHAV3D-TP small-scale phenotyping module (Figure 1d, 4a). This module utilizes the detection of TME components during IVM to offer insights into their abundance and spatial relationship with DMG cells at a single-cell level. Together with the heterogeneity module that profiles the single-cell behavior of tumor cells (Figure 1), this additional module allows linking different DMG behavioral profiles with DMG cell small-scale niches (Figures 1d, 4a). For in vivo TME labeling, we injected mice with an anti-CD31-AF647 antibody to label blood vessels, and with either SR101, a fluorescent marker known to label oligodendrocytes and astrocytes 50–53, or CD20r, a novel molecule described to label microglial cells 54 (Figure supplement 6a, b). This allowed us to measure spatial TME characteristics of each single DMG cell, such as minimal distance to a certain TME component, the number of these cell types in a certain radius, or the cell density for this particular cell (Figure 1d). Interestingly, the BEHAV3D-TP small-scale phenotyping module revealed substantial TME heterogeneity across the different behavioral DMG clusters (Figure 4b). Compared to slow and slow-retreating cells, invading cells were usually found to be less densely positioned relative to other DMG cells (Figure 4b, and Figure supplement 6c), correlating with their invasive potential into the brain parenchyma. In line with the data from the large-scale TME phenotyping (Figure 3), our data suggests that invading DMG cells reside in a local environment that is enriched in TAMMs (Figure 4b). Indeed, invading cells are closer to TAMMs compared to erratic and slow migrating cells (Figure 4c), and move directionally towards TAMMs (Figure 4d, Supplementary video 3), suggesting that migration of this behavioral cluster is at least partially mediated by chemotaxis. Furthermore, all fast-migrating behavioral clusters (invading, erratic, and retreating) were found to be generally closer to blood vessels (Figure 4b and e), suggesting that perivascular invasion is a more efficient movement route for these cells (Figure 4f), a mode of migration that has been previously observed for adult glioblastoma cells as well 7,45,47–49. Lastly, particularly slow-retreating DMG cells were found in regions enriched for SR101+ cells (Figure 4b and Figure supplement 6d), perhaps reflecting oligodendrocyte-like (OC-like) tumor DMG cells that have been found near oligodendrocytes 55. Altogether, these data demonstrate the power of our BEHAV3D-TP pipeline in correlating single-cell tumor cell behavior to its local environment, with an important opportunity to identify novel environmental regulators of invasion.

Revised version: For in vivo TME labeling, we injected mice with an anti-CD31-AF647 antibody to label blood vessels, and with either SR101, a fluorescent marker known to label oligodendrocytes and astrocytes 50–53, or CD20r, a novel molecule described to label microglial cells 54 (Figure supplement 6a, b). This allowed us to measure spatial TME characteristics of each single DMG cell, such as minimal distance to a certain TME component, the number of these cell types in a certain radius, or the cell density for this particular cell (Figure 1d). We applied option 2 of the BEHAV3D-TP small-scale phenotyping module to explore the local TME context of distinct behavioral DMG clusters, revealing substantial microenvironmental heterogeneity across clusters (Figure 4b). Compared to static cells, invading cells tended to be more sparsely distributed relative to neighboring DMG cells (Figure 4b, Figure Supplement 6c), consistent with their invasive behavior into the brain parenchyma. In agreement with our large-scale TME phenotyping results (Figure 3), invading DMG cells were found to reside in regions enriched with TAMMs (Figure 4b). Specifically, a higher proportion of invading cells were in close proximity to TAMMs compared to static, retreating, and slow-retreating clusters (Figure 4c). They frequently migrated directionally toward TAMMs (Figure 4d, Supplementary Video 3), suggesting that migration of this behavioral cluster is at least partially mediated by chemotaxis. Furthermore, all fast-migrating behavioral clusters (invading, erratic, and retreating) were found to be generally closer to blood vessels (Figure 4b and e), suggesting that perivascular invasion is a more efficient movement route for these cells (Figure 4f), a mode of migration that has been previously observed for adult glioblastoma cells as well 7,45,47–49. Lastly, particularly slow-retreating DMG cells were found in regions enriched for SR101+ cells (Figure 4b and Figure supplement 6d), perhaps reflecting oligodendrocyte-like (OC-like) tumor DMG cells that have been found near oligodendrocytes 55. Altogether, these data demonstrate the power of our BEHAV3D-TP pipeline in correlating single-cell tumor cell behavior to its local environment, with an important opportunity to identify novel environmental regulators of invasion.

Discussion

Here, we provide a new user-friendly platform tailored for IVM single tumor cell dynamics analysis, termed BEHAV3D-TP, aiming to better understand the mechanisms underlying tumor heterogeneity. In contrast to traditional pipelines that rely on the analysis of single parameters such as cell speed or displacement56–58 or that use an artificial threshold to assign a specific behavior to cells10, BEHAV3D-TP is an unbiased approach based on the integration of multiple parameters measured by IVM. Using the heterogeneity module of our BEHAV3D-TP pipeline, we uncovered behavioral phenotypes that have not been described before and that could not have been identified using single parameter analysis. For example, our analysis revealed invading and retreating clusters, composed of cells with a similar cell speed but with an opposite migration direction, an important biological and clinical difference. Likely, bidirectional migration is mediated by different chemotactic cues and by the composition of the local TME59. It is of vital importance for DMG and other cancer patients to better understand the biology underlying the different behavioral phenotypes, especially concerning invading cells that are difficult to therapeutically target 60.

Previous studies suggest a major role for the microenvironment in promoting tumor cell spread into the healthy brain 61–63. Indeed, using the complementary large- and small-scale TME phenotyping modules of BEHAV3D-TP, we identified substantial differences in the TME composition of the different behavioral DMG clusters. Our results also show that the small-scale TME analysis has a higher predictive value compared to the large-scale TME analysis, suggesting that a cell’s direct microenvironment is a more important regulator of its behavior compared to its macroenvironment. Future experiments directed towards identifying more TME components, for example using multispectral LSR-3D imaging42, could therefore further refine the large-scale phenotyping module by detecting subregion heterogeneity.

IVM enables simultaneous labeling of various TME components and allows real-time investigation of their association with tumor cell behavior. Indeed, we found that TAMMs and blood vessels were associated with specific behavioral phenotypes, showing the power of BEHAV3D-TP small-scale phenotyping module in identifying novel regulators of cellular behavior. This is exemplified by our finding that slow-retreating DMG cells were found in environments enriched for SR101+ cells, an observation that has not been made before. To take further advantage of BEHAV3D-TP, future work should be directed toward identifying more TME components, including cells of the adaptive immune system. For example, it would be interesting to use our pipeline on recently developed immunocompetent mouse models of DMG64,65, or in mice with a humanized immune system66. While our findings suggest that microenvironmental factors may influence tumor cell migration, further studies will be necessary to establish causal relationships. Additional experimental validation, such as macrophage ablation experiments, could help clarify the specific contributions of these factors.

While our current study does not provide direct functional validation of the distinct motility clusters identified, existing literature strongly supports the notion that cell dynamics can serve as a proxy for functional states and phenotypic heterogeneity. Prior work, including studies by our group19,67 as well as Crainiciuc et al.36 and Freckmann et al.20, has demonstrated that variations in cell motility patterns can reflect underlying functional characteristics. Specifically, cell morpho-dynamic features have been shown to correlate with differences in cell type identity, T-cell engagement, metastatic potential, and drug resistance states. This growing body of evidence suggests that tumor cell behavior, as captured by BEHAV3D-TP, may serve as a predictive tool for deciphering functional tumor heterogeneity. Future studies integrating transcriptomic or proteomic profiling of motility-defined subpopulations could further elucidate the biological significance of these behavioral phenotypes.

Earlier version: In addition to motility-based classification, features such as tumor cell morphology, proliferation state, and interactions with the tumor microenvironment can further refine tumor phenotyping. BEHAV3D-TP allows for the selection of diverse feature types, supporting datasets that include both dynamic, morphological and spatial parameters. However, we recognize that expanding the feature set may introduce biologically irrelevant noise, particularly in 3D microscopy data where limited z-axis resolution can lead to morphological artifacts. This highlights the potential need in the future to include unbiased feature selection strategies, such as bootstrapping methods68, to ensure the identification of meaningful and biologically relevant parameters. Careful consideration of these aspects is key to maximizing the interpretability and predictive value of analyses performed with BEHAV3D-TP.

Revised version: In addition to motility-based classification, features such as tumor cell morphology, proliferation state, and interactions with the TME can further refine tumor phenotyping. BEHAV3D-TP allows for the selection of diverse feature types, supporting datasets that include both dynamic, morphological, spatial and environmental parameters. However, we recognize that expanding the feature set may introduce biologically irrelevant noise, particularly in 3D microscopy data where limited z-axis resolution can lead to morphological artifacts. This highlights the potential need in the future to include unbiased feature selection strategies, such as bootstrapping methods68, to ensure the identification of meaningful and biologically relevant parameters. Careful consideration of these aspects is key to maximizing the interpretability and predictive value of analyses performed with BEHAV3D-TP.

To enhance the accessibility of BEHAV3D-TP for researchers keen on studying tumor cell behavioral profiles through live imaging, we have incorporated our software in a Google Colab notebook, allowing researchers with minimal programming skills to employ the tool online without the requirement to install specific software environments nor the need to directly interact with the code. The user can easily install the necessary dependencies, set up the working environment, process datasets and adjust parameters without requiring deep coding skills. Google Colab provides the necessary computational resources to be able to run the pipeline without demanding a powerful local machine. However, there are RAM usage and runtime limitations that could affect the analysis, given a big enough dataset in the free version of Google Colab. Given this and data privacy concerns, the workflow can also be run locally via a Jupyter Notebook, that can be downloaded from the Google Colab notebook. Aligned with the goals of BEHAV3D-TP, other methods utilizing the Google Colab framework are emerging, such as CellTracksColab29, which is a general tool for tracking data analysis. While CellTracksColab offers valuable insights into tracking data, BEHAV3D-TP is specifically designed for IVM data and includes additional features. These features include directionality metrics crucial for invasion studies, detailed analysis of identified behavioral populations, data back-projection, and correlation of these behaviors with both large-scale and small-scale TME features. This comprehensive approach provides a deeper understanding of the TME drivers behind intratumoral heterogeneity, making BEHAV3D-TP a robust tool for researchers investigating the complex behaviors of tumor cells in their natural environments. BEHAV3D-TP supports a broad range of data formats, including those generated by commercial Imaris and by open-source Fiji plugins such as TrackMate, MTrackJ, and ManualTracking, increasing the tool’s applicability to researchers using both commercial and open-source pipelines.

In summary, BEHAV3D-TP allows to get insights into the heterogeneity of tumor cell migration behavior in an accessible way; however, it would be exciting to extend the application to other behaviors, such as cell division or therapy response to study TME-mediated drug resistance.

Methods

Animal material

NOD.Cg-PrkdcscidIl2rgtm1Wjl/SzJ (NSG) mice were purchased from Charles River Laboratories (France). Experiments were conducted per Institutional Guidelines under acquired permission from the local Ethical Committee and as per current Dutch laws on Animal Experimentation. Mice were housed under sterile conditions using an individually ventilated cage (IVC) system and were fed with sterile food and water.

Human DMG sphere-forming culture

DMG007 (HSJD-DIPG-007) cells established from primary DMG material were grown in a base medium (TSM) consisting of 48% DMEM/F12 and 48% Neurobasal medium (Thermo Fisher) supplemented with GlutaMAX, 100mM Sodium Pyruvate, MEM non-essential amino acids and 1M Hepes buffer, at 1% each. Primocin (InvivoGen) at 50mg/ml was additionally added to the TSM. Working medium was prepared by adding 20ng/ml of human EGF and human basic FGF, 10ng/ml of PDGF-AA and PDGF-BB (Peprotech) and 2ug/ml of Heparin (StemCell Technologies). DMG cells were cultured at 37 °C in a humidified atmosphere with 5% CO2.

Fluorescent reporter cloning and lentiviral transduction

Fluorescent reporter construct was based on the combination of Histone2B-mNeonGreen, P2A and mScarlet I-CAAX gifted from Dr. Hugo Snippert (UMC Utrecht, The Netherlands), and inserted by In-Fusion HD Cloning Plus (Takarabio) into a lentiviral vector including cPPT/CTS, WPRE, an EF1a promoter and IRES-puromycin-resistance cassette11. Subsequently the above-mentioned elements were amplified by PCR and the amplicon was assembled in the final lentiviral backbone using Gibson Assembly (NEB). Sequences were validated using Sanger sequencing. DMG007 cells were transduced using a standard lentiviral transduction protocol and selected using puromycin to achieve a stable cell line.

Cranial imaging window (CIW) surgery and tumor cell injection

CIW surgery and tumor cell injection were performed on the same day as previously described 7,10. Briefly, mice were sedated with Hypnorm (Fluanison [neuroleptic] + Fentanyl [opioid]) (0.4 ml/kg) + Midazolam [benzodiazepine sedative] (2 mg/kg) at a dose of 1:1:2 in sterile water and mounted in a stereotactic frame. The head was secured using a nose clamp and two custom-made 3D printed ear bars, made out of Polylactic Acid (PLA, Geeetech, 700-001-043). The head was shaved, and the skin was cut circularly. The mouse was placed under a stereo-microscope to ensure precise surgical manipulation. The periosteum was scraped, and a circular groove of 5 mm diameter was drilled over the right parietal bone. The bone flap was lifted under a drop of cortex buffer (125 mM NaCl, 5 mM KCl, 10 mM glucose, 10 mM HEPES buffer, 2 mM MgSO4, and 2 mM CaCl2, pH 7.4), and the dura mater was removed. Gelfoam sponge (Pfizer) was used to stop potential bleedings. Next, 1 × 105 DMG007-H2B-mNeonGreen cells resuspended in 3 μl phosphate-buffered saline (PBS) were injected stereotactically using a 5 μl Hamilton syringe with a 2 pt style in the middle of the craniotomy at a depth of 0.5 mm. The exposed brain was sealed with silicone oil and a 6 mm coverslip was glued on top. Dental acrylic cement (Vertex) was applied on the skull surface, covering the edge of the coverslip, and a custom-made 3D printed ring made of biocompatible Polylactide acid (PLA, Geeetech, 700-001-0433) was glued around the coverslip to provide fixation during intravital imaging and to serve as a reservoir for a water drop required for the objective of the microscope. A single dose of 100 μg/kg of buprenorphine (Temgesic, BD pharmaceutical limited) was administered for post-surgical pain relief. Mice were closely monitored twice per week to assess behavior, reactivity, and appearance.

Intravital imaging

For DMG IVM imaging mice were intravitally imaged as previously described 7,10. In short, mice were sedated with isoflurane and placed face-down in a stereotactic frame. The Polylactide ring was used to fix the mouse’s head to the frame with custom-made Polylactide bars. Time-lapse images of the entire tumor volume were acquired for a maximum of 5.4 hours. The minimal time interval between serial images was set to 20 minutes. For tile scans, images of the complete z-stack of the tumor were acquired, with a step size of 2 µm. In a group of 3 mice, oligodendrocytes and astrocytes were imaged by an intravenous injection of 50 µL SR101 (Thermofisher) at a concentration of 5 mM dissolved in PBS. In a different group of 3 mice, TAMMs were imaged by intravenous injection of 50 μL 100 µM CDr2054. Simultaneously with the injection of either SR101 or CDr20, all mice received an intravenous injection of 10 ul CD31-AF647 antibody (ThermoFisher, A14716) to visualize the blood vessels.

Intravital imaging was performed on an upright FVMPE-RS multiphoton microscope (Olympus) equipped with a MaiTai DeepSee and an InSight DeepSee laser (Spectra-Physics), a 25x/1.05 numerical aperture water objective, and two GaAsP and two Multialkali photomultiplier tubes (Olympus). Images were acquired with simultaneous imaging with the MaiTai (960 nm) and Insight (1100 nm) lasers. Green, red, and far-red fluorescence was separated by 570 nm and 650 nm dichroic mirrors, and 610/70 nm and 705/90 nm filters (Semrock) and collected by GaAsP or Multialkali photomultiplier tubes. Scanning was performed in a bidirectional mode with a resonant scanner at 400 Hz and 12 bit, with a zoom of 1x, 512 × 512 pixels. For optimal detection of blood vessels (CD31-AF647), a single tile scan was made at the last timepoint with Insight laser tuned at 1250 nm. At the postprocessing step, this image was co-registered to the time-lapse movie.

Breast tumor IVM imaging was performed as described before15 and data was kindly provided by Dr. Nienke Vrisekoop. GBM IVM imaging was performed as described before7 and data was kindly provided by Prof. Jacco van Rheenen. Breast epithelial IVM imaging and cell tracking with MaSCOT-AI was performed as described before31 and single cell tracks were downloaded from Zenodo repository (14503476).

Co-registration of IVM imaged blood vessel imaging

CD31-AF647 tile scan featuring optimal visualization of blood vessels during live imaging was co-registered to the last timepoint of the time-lapse movie. Both images (last timepoint of time-lapse tile-scan and blood vessel tile-scan) were visually inspected to determine and annotate the internal landmarks using Imaris (Oxford Instruments), versions 9.5 to 9.6. These landmarks served as a reference for the co-registration software to overlap both images. To perform the co-registration, a modification of elastix69 was used; a python-based open-source software based on Insight Segmentation and Registration Toolkit (ITK) plus SimpleElastix70 : a collection of different algorithms used on medical image registration. Due to the extensive number of potential parameter combinations, specific parameters were manually chosen. This selection was conducted using the Slicer3D visualization software, version 5.2.2.

Immunohistochemistry

3D immunohistochemistry was performed as described before 43. After the live imaging session, mouse brains were immersed in 5 ml 4% paraformaldehyde (PFA) pH 7.4 overnight on ice. After fixation, brains were blocked in 5-10 ml Wash Buffer 1 (2 ml Tween-20, 2 ml Triton X-100, 2 ml 10% SDS, and 2 g BSA in 1 l PBS) for 5 hours. For multiplex immunolabeling, a two- or three-round staining protocol was used as previously described42,43. Washing and incubation steps were performed in Wash Buffer 2 (1 ml Triton X-100, 2 ml 10% SDS, and 2 g BSA in 1 l PBS). A thick (1 mm) slice of the cortex part that was under the cranial imaging window was dissected with a scalpel before proceeding to immunolabeling. The following primary antibodies were used: rabbit anti-Olig2 1:100 (AB9610, Merk), goat anti-Iba1 1:200 (NB100-1028SS, NovusBio) and CD31-AF647 1:250 (A14716, ThermoFisher). As secondary antibodies, donkey anti-rabbit AF405 1:250 (ab175651, Abcam) and donkey anti-goat AF633 1:250 (A21082, ThermoFisher) were used. After the last washing step, tissues were optically cleared by three stepwise incubations of an increasing concentration of FUnGI71 clearing agent (25%/50%/75%, diluted in PBS) for 1h at room temperature42. The final incubation step with 100% FUnGI was performed overnight at 4°C.

mLSR3D imaging

mLSR3D imaging42,43 of thick slices of fixed mouse cortex was performed using a Zeiss LSM880 system equipped with a 32-channel spectral detector. The signal from all fluorophores present in the samples was collected in a single acquisition and linear unmixing was used to obtain separate channels. Unmixing was carried out by the Online Fingerprinting mode of the microscope. Pre-acquired references of each single fluorophore were used to enable multispectral on-the-fly unmixing. Imaging was performed using a 25x multi-immersion objective 0.8 NA (Zeiss 420852-9871) with a working distance of 570µm. Tile scans were acquired using a voxel size of 0.332×0.332×1.200 µm, a pixel dwell of 2µs, and a 10% tile overlap.

Imaris image processing

DMG IVM datasets were processed using Imaris software (Oxford Instruments, versions 9.5 to 10.1) for 3D visualization, shift correction, object rendering, and cell tracking of time-lapse movies.

Imaris Cell tracking

The Channel Arithmetics Xtension was used for channel unmixing of overlapping signals. The Spots ImarisTrack module was used for object detection and semi-automated tracking of tumor cells (autoregressive motion). Tumor cell tracks were manually corrected to ensure accurate tracking. Around 50-100 cells per position were manually corrected and labeled as such for posterior selection of accurate tracks. To determine the direction of tumor cell movement relative to the tumor core, at the image edge closer to the tumor core a “Tumor edge” surface object was manually rendered using the contour tool. The Distance Transformation Xtension was used to measure the distance between tumor cells and the “Tumor edge”. For tracked tumor cells, time-lapse data containing the coordinates of each cell, the values of cell speed, mean square displacement, displacement delta length, displacement length, and distance to “Tumor edge” were exported. See https://qbi.uq.edu.au/files/40820/ImarisManual.pdf for statistical details.

Imaris Microenvironmental factors rendering

For rendering the microenvironmental factors detected during IVM or by post-IVM LSR-3D imaging, the Surface and Spots modules of Imaris were used. Surfaces were used for the SR101, CDr20, Olig2, DAPI, and Iba1 stainings, and Spots were used for CD31+ protruding or elongated structures, as previously described 44. For rendered structures, positional data containing the coordinates of objects were exported for large-scale spatial phenotyping with Cytomap.

Fiji image processing

μSAM cell segmentation

Segmentation was performed using the μSAM plugin for Napari33, where vit-b model was selected to compute embeddings for each timeframe. Automatic segmentation was applied, and parameters were adjusted for optimal results. Manual corrections were made to remove unwanted surfaces or add missing cells. Once segmentation was complete, frames were combined into a single stack in ImageJ and tracked with Trackmate.

Fiji tracking

Breast tumor IVM data was manually tracked with Fiji MtrackJ plugin as previously described10,15.

Adult glioma cells, segmented with μSAM, were consequently tracked with Fiji TrackMate plugin. The labeled image stack, segmented using μSAM, was imported into ImageJ for cell tracking with the Fiji TrackMate plugin. The “Label Image Detector” was employed to identify individual cells, and the LAP Tracker was used to link cells across timeframes. Tracking parameters were optimized to account for missing frames and ensure accurate cell trajectories. The statistics for each tracked cell were subsequently exported. To integrate both dynamic and morphological data, a custom script was developed to merge the dynamic information from TrackMate with the morphological data obtained from the μSAM segmentation labels for each individual cell. This integration was achieved by assigning the “TRACK_ID” from the TrackMate “allspots” output file to each segmented instance extracted from the TIFF file. The assignment was performed using the “measure” function in scikit-image and the KDTree distance function from the scipy module, based on the x, y, and z positions of the cells.

The script is available both as a downloadable version on GitHub (https://github.com/imAIgene-Dream3D/MorphoTrack-merger) and as an online application (https://morphotrack-merger.streamlit.app/), making this tool more accessible for non-coding users.

Tumor large-scale spatial phenotyping with Cytomap

The positional data of various environmental surface features, segmented from LSR-3D imaging data, was analyzed for large-scale spatial phenotyping using the Cytomap Spatial Analysis Toolbox44. Briefly, neighborhoods with a 100 µm radius were generated. Neighborhoods were then clustered into 2 regions based on DMG cell positioning using the David Balwin algorithm. Using the gating tool, the tumor regions (containing DMG objects) were defined. Next, the neighborhoods were clustered again into 3 different regions based on the distribution of Iba1, Olig2, and CD31: Void, TAMM/Oligo, and TAMM/vascularized. By using the Manually Defined Regions option, only the tumor region was classified. Finally, for each region, fold-change values were calculated and exported using the Region Statistics tab. To map the assigned regions onto IVM movies, a 3D image of the cluster distribution within the tumor was generated and exported for each sample (Figure Supplement 5a). Next, regions within the IVM movies were visually matched to the corresponding regions identified by the Large-Scale Phenotyping module of Cytomap (Figure 3c). For each mouse, at least one or two representative positions per matched region type were selected, cropped, and analyzed to assess tumor cell behavior, following the previously described cell tracking methodology (Imaris Cell tracking).

BEHAV3D Tumor Profiler framework

BEHAV3D-TP was developed using the R Studio version 4.3.3. It features three modules: 1) Heterogeneity module with optional 2) Large-scale and/or 3) Small-scale phenotyping module. These modules are described in detail below, exemplified by the analysis of DMG data from this study. All of the modules are included in the same script available in Github [https://github.com/imAIgene-Dream3D/BEHAV3D_Tumor_Profiler], but can be run independently from one another. To acquire this independency, the optional modules are directly combined to the necessary sections of the Heterogeneity module. Consequently, the user can use whichever module is best suited for their purposes. The complete workflow— combining heterogeneity analysis with large- and small-scale phenotyping—is primarily optimized for datasets processed with Imaris due to its wide use in the IVM community and special suitability for 3D visualization and tracking. In order to further facilitate public access to the BEHAV3D-TP, a Google Colab notebook has also been created accessible through the same Github page [https://github.com/imAIgene-Dream3D/BEHAV3D_Tumor_Profiler]. This environment is friendly to anyone regardless of their coding expertise as it has a step-by-step follow-through execution through the pipeline.

In addition, to broaden accessibility and support alternative data processing pipelines, the core Heterogeneity module of BEHAV3D-TP is also compatible with Fiji-based tracking data, specifically outputs generated by TrackMate. MTrackJ and ManualTracking. A separate Google Colab notebook is provided for these workflows on the Github page, allowing users to easily import Fiji-processed datasets, adjust analysis parameters, identify different morpho-dynamic single cell patterns.

A Wiki demo run and a video tutorial (Supplementary Video 1) are also provided through the GitHub page [https://github.com/imAIgene-Dream3D/BEHAV3D_Tumor_Profiler] for further clarification.

To ensure that your input data aligns with the workflow, we provide an online application for data quality control (https://github.com/alievakrash/BEHAV3D_TP_dataQC). This tool allows users to visualize various features, grouped by a condition of interest, and check for issues such as a high number of missing values (NA). It enables users to assess their data for any inconsistencies or unexpected patterns before proceeding with further analysis.

Heterogeneity module

To account for missing time points and to create a time series with the same time interval for each tumor cell time series, linear interpolation was used to estimate the values of missing time points. To compare time series independently, tracks were cropped to the minimal common length among all tracks. In the case of DMG data to a 2.6 hour duration. For tracked tumor cells, time-lapse data containing the coordinates of each cell, the values of cell speed, mean square displacement, displacement delta length, displacement length, and relative distance to “Tumor edge” were used to calculate a cross-distance matrix between time series. First, we performed a principal component (PC) analysis and selected the number of PCs that explained at least 90% of data variance (3).

The previously calculated PCs were then used to compute the cross-distance matrix between each tumor cell multivariate time series using the dynamic time warping algorithm from the package “dtwclust”. To visualize heterogenous tumor cell behaviors in 2 dimensions, dimensionality reduction on the cross-distance matrix was performed by the Uniform Manifold Approximation and Projection method (“umap” package). Clustering was performed using the k-means clustering algorithm. For each cluster, the average values of distinct behavioral and environmental components were calculated and plotted as a heatmap (“pheatmap”).

Large-scale phenotyping module

The information obtained from mLSR3D and Cytomap about large-scale TME regions were combined with the DMG cells behavioral information (see above section Tumor large-scale spatial phenotyping with Cytomap) to examine the large-scale features of tumor cells. For tumor cells, the mean speed, square displacement and raw movement were analyzed.

Small-scale phenotyping module

Added in the revised version: BEHAV3D-TP enables the integration of microenvironmental features—beyond tumor cells— through two distinct approaches within this module (Figure 1d). In Option 1, these microenvironmental features (MEFs) are incorporated upstream and directly included in the clustering process. Feature selection for this step is performed via the interface titled “Select the features you want to use for the environmental (morpho)-dynamic analysis.” In Option 2, illustrated in Figure 4, environmental features are integrated downstream to assess their correlation with pre-defined behavioral clusters. In this case, features are selected through a separate interface titled “Feature selection for projection over UMAP.”

Earlier version: The positional coordinates of tracked tumor cells and detected microenvironmental factors were used to determine the small-scale landscape of tumor cells. For tumor cells, the average distance to the closest 10 neighbors was measured to determine tumor cell density. For each tumor cell and each environmental component, the number of microenvironmental objects in a 30-µm radius and the distance to the closest object were computed. For each behavioral cluster, the values of closest neighbors, minimal distance to each environmental factor, and number of environmental objects in the 30 µm radius were plotted and compared.

Revised version: In the DMG example of Figure 4 the positional coordinates of tracked tumor cells and detected microenvironmental factors (TAMMs, SR101+ and blood vessels) were used to determine the small-scale landscape of tumor cells. For tumor cells, the average distance to the closest 10 neighbors was measured to determine tumor cell density. For each tumor cell and each environmental component, the number of microenvironmental objects in a 30-µm radius and the distance to the closest object were computed. For each behavioral cluster, the values of closest neighbors, minimal distance to each environmental factor, and number of environmental objects in the 30 µm radius were plotted and compared. To quantify the percentage of DMG cells in proximity to CD20r⁺ TAMMs, we classified cells located within 15 µm of TAMM surfaces as “close”. For proximity to blood vessels, DMG cells were considered “close” if they were within 3 µm of blood vessel spots. The difference in distance thresholds reflects the nature of the reference objects: while TAMMs are defined by “Surface” boundaries, blood vessels were represented by smaller centroid-based “Spots”, necessitating a more stringent cutoff.

Statistical analyses

Earlier version: Statistical analyses were performed using R. The representation of n may refer to individual mice, distinct imaging positions within different mice, or individual cell tracks, as indicated in the figure legends. Two-tailed unpaired t-tests were performed for comparison between two groups unless indicated otherwise (Figure supplement 1d). For comparisons of more than two groups we performed a one-way ANOVA (Figure 2d and Figure 4b) followed by Tukey HSD correction (Figure Supplement 4c (p values in legend), Figure 3d, Figure 4 c,e, Figure Supplement b, e-g, Figure Supplement 6 c,d). To compare the distribution of different morphodynamic clusters of TEB cells, we pooled all positions and mice due to the low number of cells per position. We then constructed a contingency table and performed a Chi-square test (Figure supplement 2d). To compare the distance to the nearest neighbors between morphodynamic clusters of GBM cells, a mixed-effects linear regression was performed with mouse as a random effect (Figure supplement 3d).

Revised version: Statistical analyses were performed using R. The representation of n may refer to individual mice, distinct imaging positions within different mice, or individual cell tracks, as indicated in the figure legends. Two-tailed unpaired t-tests were performed for comparison between two groups in Figure supplement 1d. For comparisons involving more than two groups, where no random effects were present, we performed one-way ANOVA (Figure 2d and Figure 4b), followed by Tukey’s Honest Significant Difference post hoc test to correct for multiple comparisons. This correction was applied in Figure Supplement 4c (p-values reported in the legend), as well as in Figure 3d, and Figure Supplement 5b, 5e–g. Where necessary to account for inter-mouse and TME region variability, we applied a linear mixed-effects model with mouse and/or TME class included as random effects (Figure 4c,e, Figure Supplement 3d, Figure Supplement 6c,d). To compare the distribution of different morphodynamic clusters of TEB cells, we pooled all positions and mice due to the low number of cells per position. We then constructed a contingency table and performed a Chi-square test (Figure supplement 2d).

Removed in the revised version: Data availability

Removed in the revised version: DMG imaging data is publicly available at BioStudies accession number S-BIAD1759 (https://www.ebi.ac.uk/biostudies/studies/S-BIAD1759). Breast cancer IVM data was previously described15 and kindly provided by Dr Nienke Vrisekoop. Adult glioma IVM data was previously described7 and kindly provided by Prof. Jacco van Rhennen. Healthy breast epithelium IVM was previously described31 and kindly provided by Prof. Jane Visvader. Demo outputs and Jupyter notebooks of Imaris- and Fiji-processed data are provided in our GitHib repository (https://github.com/imAIgene-Dream3D/BEHAV3D_Tumor_Profiler).

Removed in the revised version: Code availability

Removed in the revised version: We provide the BEHAV3D-TP framework in a Google Colab user-friendly environment [https://colab.research.google.com/drive/1JI7ysqFf3tvdi6Df4YUsSZ8RbuXw8wba?usp=shari ng], as well as in an R script format in Github [https://github.com/imAIgene-Dream3D/BEHAV3D_Tumor_Profiler]. Additionally, a Wiki [https://github.com/imAIgene-Dream3D/BEHAV3D_Tumor_Profiler/wiki] to follow through a demo dataset as well as a video tutorial (Supplementary Video 1) has been developed for further assistance. For the BEHAV3D-TP heterogeneity module, we offer an online app to assess data quality (https://behav3d-tp-data-qc.streamlit.app/), with the corresponding code available on GitHub (https://github.com/alievakrash/BEHAV3D_TP_dataQC). Additionally, for quantifying morphological features derived from Microsam segmentation and integrating them with dynamic features from TrackMate, we provide another online app (https://morphotrack-merger.streamlit.app/) and related code on GitHub (https://github.com/imAIgene-Dream3D/MorphoTrack-merger).